---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Lesson 2.1 - Introduction to logistic regression modeling

```yaml
type: VideoExercise
key: ff9a389130
xp: 50
```

`@projector_key`
3e2d624300abd9e269bcd4230b0a5212
